# Deep Learning

In this exercise, you will use a deep neural network to predict the values of houses based on some provided input data. You will use keras to build the model. Below is a description of how the keras models are set up.

```python
import tensorflow as tf
import tensorflow.keras as keras
import numpy as np

np.random.seed(0)
```

```python
(x_train, y_train), (x_test, y_test) = tf.keras.datasets.boston_housing.load_data()
```

Output:
```
Downloading data from https://storage.googleapis.com/tensorflow/tf-keras-datasets/boston_housing.npz
57344/57026 [==============================] - 0s 0us/step
```

```python
x_train.shape
```

Output:
```
(404, 13)
```

```python
y_train.shape
```

Output:
```
(404,)
```

```python
y_train.mean(), y_train.std()
```

Output:
```
(22.395049504950492, 9.199035423364862)
```

The keras model consists of multiple parts:

1. Construct the model layers, neurons per layer, activation function
1. Determine the loss function, metrics and optimization method
1. Fit the model to some data
1. Evaluate the model using the same metric

Some relevant docs:
 - [initializers](https://keras.io/initializers/)
 - [loss functions](https://www.tensorflow.org/api_docs/python/tf/keras/losses)
 - [regularizations](https://keras.io/regularizers/)
 - [optimizers](https://keras.io/optimizers/)
 - [metrics](https://www.tensorflow.org/api_docs/python/tf/keras/metrics)


First, to construct a model, use the [Sequential](https://keras.io/getting-started/sequential-model-guide/) object. You can pass multiple layers to the sequential object. For this exercise, we will only be using the [Dense](https://keras.io/layers/core/#dense) layers.

<!-- grade_id: cell-51879de7032d8823, points: 0, grade: false, solution: true -->
```python
# Create 3 hidden layers and an output layer with 13, 6, 3 neurons respectively and set them to layers
# Use any activation you like such as "relu" or "tanh", you can alternate for each layer
# YOUR CODE HERE
#raise NotImplementedError()

layers = [
    keras.layers.Dense(13, input_shape=(13,), activation='relu'), 
    keras.layers.Dense(6, activation='tanh'), 
    keras.layers.Dense(3, activation='relu'), 
    keras.layers.Dense(1, activation='tanh')
]
model = keras.Sequential(layers)

```

```python
len(layers)
```

Output:
```
4
```

<!-- grade_id: cell-ee74eff075f222ed, points: 5, grade: true, solution: false -->
```python
assert len(layers) == 4
for i,layer in enumerate(layers):
    assert isinstance(layers[i], keras.layers.Dense) 
    assert layer.weights[1].shape == [13,6,3,1][i]
```

<!-- grade_id: cell-9c5dcca57115d02a, points: 0, grade: false, solution: true -->
```python
# Set up the model to do the following:
# - use stochastic gradient descent to fit the model
# - use mean absolute error as its loss function
# - use mean absolute error as one of its metrics

# YOUR CODE HERE
model.compile(optimizer='SGD',
              loss='mean_absolute_error',
              metrics=['mean_absolute_error'])

```

<!-- grade_id: cell-0746e29c34118826, points: 5, grade: true, solution: false -->
```python
assert isinstance(model.optimizer, keras.optimizers.SGD)
assert model.loss in ["mae", "mean_absolute_error"]
assert "mae" in model.metrics or "mean_absolute_error" in model.metrics
```

```python
# Now fit the model and print a summary for it
model.fit(x_train, y_train, epochs=50)
model.summary()
```

Output:
```
Epoch 1/50
404/404 [==============================] - 0s 1ms/step - loss: 22.3371 - mean_absolute_error: 22.3371
Epoch 2/50
404/404 [==============================] - 0s 41us/step - loss: 22.2114 - mean_absolute_error: 22.2114
Epoch 3/50
404/404 [==============================] - 0s 42us/step - loss: 22.0964 - mean_absolute_error: 22.0964
Epoch 4/50
404/404 [==============================] - 0s 42us/step - loss: 21.9963 - mean_absolute_error: 21.9963
Epoch 5/50
404/404 [==============================] - 0s 42us/step - loss: 21.9120 - mean_absolute_error: 21.9120
Epoch 6/50
404/404 [==============================] - 0s 41us/step - loss: 21.8424 - mean_absolute_error: 21.8424
Epoch 7/50
404/404 [==============================] - 0s 43us/step - loss: 21.7853 - mean_absolute_error: 21.7853
Epoch 8/50
404/404 [==============================] - 0s 43us/step - loss: 21.7386 - mean_absolute_error: 21.7386
Epoch 9/50
404/404 [==============================] - 0s 47us/step - loss: 21.7002 - mean
```

```python
# Here we can evaluate how our model does based on the test data
model.evaluate(x_test, y_test)
```

Output:
```
102/102 [==============================] - 0s 818us/step
[22.122359518911324, 22.122359518911324]
```

Now let's try another optimizer instead of stochastic gradient descent. [Adam](https://keras.io/optimizers/#adam) is the recommended default for training neural networks since it usually performs quite well. In the next cell, compile the model with adam instead of sgd and use the same loss and metrics. After compiling, fit the data for as many epochs as you think it takes to see the value start to converge.

<!-- grade_id: cell-bee688f3d919be61, points: 0, grade: false, solution: true -->
```python
# Compile the model using adam
# YOUR CODE HERE

model.compile(optimizer='Adam',
              loss='mean_absolute_error',
              metrics=['mean_absolute_error'])
```

<!-- grade_id: cell-fe8cd91b8b780964, points: 5, grade: true, solution: false -->
```python
assert isinstance(model.optimizer, keras.optimizers.Adam)
assert model.loss in ["mae", "mean_absolute_error"]
assert "mae" in model.metrics or "mean_absolute_error" in model.metrics
```

```python
model.evaluate(x_test, y_test)
```

Output:
```
102/102 [==============================] - 0s 2ms/step
[22.122359518911324, 22.122359518911324]
```

Now recreate the model with dropout layers. Add 2 dropout layers, each after the first or second layer respectively. Select a low value of dropout that results in a good score.

<!-- grade_id: cell-d5040e44398f8dc9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
layers = [
    keras.layers.Dense(13, input_shape=(13,), activation='relu'), 
    keras.layers.Dropout(0.1),
    keras.layers.Dense(6, activation='tanh'),
    keras.layers.Dropout(0.01),
    keras.layers.Dense(3, activation='relu'), 
    keras.layers.Dense(1, activation='tanh')
]
model = keras.Sequential(layers)


model.compile(optimizer='Adam',
              loss='mean_absolute_error',
              metrics=['mean_absolute_error'])
```

<!-- grade_id: cell-b6c7e6e2ddc6742a, points: 5, grade: true, solution: false -->
```python
assert len(layers) == 6
assert isinstance(layers[1], keras.layers.Dropout)
assert isinstance(layers[3], keras.layers.Dropout)
for i,layer in enumerate(layers):
    if i not in [1,3]:
        assert isinstance(layers[i], keras.layers.Dense) 
        assert layer.weights[1].shape == [13,0,6,0,3,1][i]
```

```python
model.compile(optimizer='adam', loss='mae', metrics=['mae', "mse"])
model.fit(x_train, y_train, epochs=500, verbose=0)
model.summary()
```

Output:
```
_________________________________________________________________
Layer (type)                 Output Shape              Param #   
=================================================================
dense_74 (Dense)             (None, 13)                182       
_________________________________________________________________
dropout_6 (Dropout)          (None, 13)                0         
_________________________________________________________________
dense_75 (Dense)             (None, 6)                 84        
_________________________________________________________________
dropout_7 (Dropout)          (None, 6)                 0         
_________________________________________________________________
dense_76 (Dense)             (None, 3)                 21        
_________________________________________________________________
dense_77 (Dense)             (None, 1)                 4         
=================================================================
Total para
```

```python
model.evaluate(x_test, y_test)
```

Output:
```
102/102 [==============================] - 0s 93us/step
[22.078435561236212, 22.078435561236212, 570.7011335784314]
```

Select a dropout rate that gets an okay score.

## Feedback

<!-- grade_id: feedback, points: 0, grade: false, solution: true -->
```python
def feedback():
    """Provide feedback on the contents of this exercise
    
    Returns:
        string
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: feedback-tests, points: 0, grade: true, solution: false -->
```python

```